---
title: Chapter 1. Regression and the Normal Distribution
description: >-
  This is a template chapter. Okay


---
## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 6154f289da
```

Do some data science. Okay

`@instructions`
Frees testy

`@hint`


`@pre_exercise_code`
```{r}
heights1 <- read.csv("datasets/GaltonFamily.csv",header=TRUE)
```
`@sample_code`
```{r}
str(heights1)
```
`@solution`
```{r}
str(heights1)
```
`@sct`
```{r}
str(heights1)
```




